```python
import os 
import pickle
from time import time
```

```python
_version = '10gb'
```

```python
# Replace information
with open('HW2.py', 'r', encoding='utf-8') as file:
    data = file.read()
    
source = "_is_in_develop = True"
output = "_is_in_develop = False"
data = data.replace(source, output, 1)

source = f"_version = '5mb'"
output = f"_version = '{_version}'"
data = data.replace(source, output, 1)

filename = f'HW2-{_version}.py'
with open(filename, 'w', encoding='utf-8') as file:
    file.write(data)
```

```python
if os.path.exists('answer.pickle'):
    !rm answer.pickle

start = time()
!timeout 1420 python {filename}
end = time()

duration = end-start

# Read answers and delete temp files
with open('answer.pickle', 'rb') as file:
    answers = pickle.load(file)
!rm answer.pickle
!rm {filename}
```

Output:
```
Setting default log level to "WARN".
To adjust logging level use sc.setLogLevel(newLevel). For SparkR, use setLogLevel(newLevel).
23/05/02 06:40:40 WARN NativeCodeLoader: Unable to load native-hadoop library for your platform... using builtin-java classes where applicable
Username = jiy068
23/05/02 06:40:58 WARN MemoryStore: Not enough space to cache rdd_1_45 in memory! (computed 2.3 MiB so far)
23/05/02 06:40:58 WARN BlockManager: Block rdd_1_45 could not be removed as it was not found on disk or in memory
23/05/02 06:40:58 WARN BlockManager: Putting block rdd_1_45 failed
23/05/02 06:40:58 WARN MemoryStore: Not enough space to cache rdd_1_44 in memory! (computed 5.2 MiB so far)
23/05/02 06:40:58 WARN BlockManager: Block rdd_1_44 could not be removed as it was not found on disk or in memory
23/05/02 06:40:58 WARN BlockManager: Putting block rdd_1_44 failed
23/05/02 06:40:58 WARN MemoryStore: Not enough space to cache rdd_1_46 in memory! (computed 2.3 MiB so far)
23/05/02 06:40:58 WARN 
```

<!-- grade_id: cell-624003ee932f5c48, points: 0, grade: true, solution: false -->
```python
##Hidden test cases here
```

<!-- grade_id: num-tweets, points: 2, grade: true, solution: false -->
```python
##Hidden test cases here
```

<!-- grade_id: num-unique-users, points: 2, grade: true, solution: false -->
```python
##Hidden test cases here
```

<!-- grade_id: counts_per_part, points: 2, grade: true, solution: false -->
```python
##Hidden test cases here
```

<!-- grade_id: num_of_tokens, points: 2, grade: true, solution: false -->
```python
##Hidden test cases here
```

<!-- grade_id: num_freq_tokens, points: 2, grade: true, solution: false -->
```python
##Hidden test cases here
```

<!-- grade_id: top20, points: 2, grade: true, solution: false -->
```python
##Hidden test cases here
```

<!-- grade_id: popular_10_in_each_group, points: 2, grade: true, solution: false -->
```python
##Hidden test cases here
```

```python
print('Your total time taken is : ',duration)
```

Output:
```
Your total time taken is :  886.3436224460602
```

<!-- grade_id: runtime_test_7mb_1, points: 6, grade: true, solution: false -->
```python
assert duration < 1315, "Your implementation is slower than the Base solution."
```

<!-- grade_id: runtime_test_7mb_2, points: 4, grade: true, solution: false -->
```python
assert duration <= 1175, "Your implementation is slower than our Available Key"
```

```python
assert duration <= 650, "Your implementation is slower than our Best Solution."
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[16], line 1
----> 1 assert duration <= 650, "Your implementation is slower than our Best Solution."

AssertionError: Your implementation is slower than our Best Solution.
```